<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 0  -->
  <question type="category">
    <category>
        <text>top</text>

    </category>
  </question>

<!-- question: 0  -->
  <question type="category">
    <category>
        <text>top/all</text>

    </category>
  </question>

<!-- question: 0  -->
  <question type="category">
    <category>
        <text>top/all/dev</text>

    </category>
  </question>

<!-- question: 817  -->
  <question type="stack">
    <name>
      <text>実習出題用:大きな転置推移確率行列を持つマルコフ連鎖</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>転置推移確率行列\(M\)を持つマルコフ連鎖について, プログラムを作成して, 分布の時間発展を調べなさい.</p><p>あなたの考える転置推移確率行列\(M\)は, 
\({@msize@}\times {@msize@}\)で, 次の通りです.</p>
\[ M= \begin{pmatrix} {@(bctype-2)*(bctype-3)*1/2+(bctype-1)*(bctype-2)*(1-acoeff-bcoeff)/2+(bctype-1)*(bctype-3)*(1-acoeff-bcoeff)*(-1)@} &amp; {@bcoeff@} &amp; 0 &amp; \cdots &amp; 0 &amp; {@(bctype-1)*(bctype-2)*acoeff/2@} \\ {@(bctype-1)*(bctype-3)*(acoeff+bcoeff)*(-1)+(bctype-1)*(bctype-2)*acoeff/2+0@} &amp; {@1-acoeff-bcoeff@} &amp; {@bcoeff@}&amp; \ddots &amp; &amp;0\\ 0 &amp; {@acoeff@}&amp; {@1-acoeff-bcoeff@} &amp; \ddots&amp;\ddots&amp;\vdots \\ \vdots &amp; \ddots &amp;\ddots &amp; \ddots&amp; {@bcoeff@} &amp; 0\\ 0 &amp; &amp; \ddots &amp; {@acoeff@} &amp;{@(1-acoeff-bcoeff)@}&amp; {@(bctype-1)*(bctype-3)*(acoeff+bcoeff)*(-1)+(bctype-1)*(bctype-2)*bcoeff/2+0@} \\ {@(bctype-1)*(bctype-2)*bcoeff/2@} &amp; 0 &amp; \cdots &amp; 0 &amp;{@acoeff@}&amp; {@(bctype-2)*(bctype-3)*1/2+(bctype-1)*(bctype-2)*(1-acoeff-bcoeff)/2+(bctype-1)*(bctype-3)*(1-bcoeff-acoeff)*(-1)@} \\ \end{pmatrix}. \]<p>&nbsp;&nbsp;</p><p>斜めの点\(\ddots\)に沿っては同じ値が並びます.</p>
<p>答は1と入力しておいてください.</p>
<p>[[input:ans1]] [[validation:ans1]]</p><p><br></p><p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <stackversion>
      <text>2018080600</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[msize:rand_with_prohib(20,40,[])
bctype:rand_with_prohib(1,3,[])
bcname:["吸収境界条件","反射境界条件","周期的境界条件"]
denom:10
alcoeff:rand_with_prohib(1,3,[])
bicoeff:rand_with_prohib(alcoeff+1,denom-alcoeff,[alcoeff])
dummy:"for the particle to move in negative direction"
acoeff:alcoeff/denom
bcoeff:bicoeff/denom

]]></text>
    </questionvariables>
    <specificfeedback format="moodle_auto_format">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote>
      <text>{@[msize,bctype,acoeff,1-acoeff-bcoeff,bcoeff]@} {@msize@} x {@msize@},  pn[x]={@acoeff@}p[x-1]+{@(1-acoeff-bcoeff)@}p[x]+{@bcoeff@}p[x+1], {@bctype@}=(1:absorbing, 2:reflecting, 3:periodic), {@bcname[bctype]@}</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="moodle_auto_format">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="moodle_auto_format">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="moodle_auto_format">
      <text></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>[msize,bcname[bctype],acoeff*p(x-1),(1-acoeff-bcoeff)*p(x),bcoeff*p(x+1)]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>